<a href="https://colab.research.google.com/github/kl20879/CE880-lad-class/blob/main/week1/problem2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-1 : Problem 2 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

---
## Question 4:

Write a Python program to print the even numbers from a given list.

Sample List : `[1, 2, 3, 4, 6, 7, 8, 9]`
Expected Result : `[2, 4, 6, 8]`

<!-- grade_id: even_num_ans, points: 0, grade: false, solution: true -->
```python
def even_num(lst_1):
    """Write a Python program to print the even numbers from a given list."""
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
even_num([2, 3, 14, 17, 8, 19, 11, 14])
```

<!-- grade_id: cell-87c42cb5c0f501eb, points: 1, grade: true, solution: false -->
```python
even_num([1, 2, 3, 4, 6, 7, 8, 9]) == [2, 4, 6, 8]
even_num([2, 3, 4, 7, 8, 9, 11, 14]) == [2, 4, 8, 14]
even_num([2, 3, 14, 17, 8, 19, 11, 14]) == [ 2,  8, 14]
```

---
## Question 5:

Write a Python function to create and print a list where the values are square of numbers between 1 and `n`, where `n` is user defined number: Example sqr_values(4) == `[1, 4, 9]`

<!-- grade_id: sqr_values_ans, points: 0, grade: false, solution: true -->
```python
def sqr_values(n):
    """Write a Python program to print the square of numbers between 1 and `n`"""
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
sqr_values(4)
```

<!-- grade_id: sqr_values_test, points: 1, grade: true, solution: false -->
```python
"""Check that sqr_values returns the correct answer for various inputs."""
assert sqr_values(4) == [1, 4, 9]
assert sqr_values(8) == [1, 4, 9, 16, 25, 36, 49]
assert sqr_values(9) == [1, 4, 9, 16, 25, 36, 49, 64]
assert sqr_values(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

## Question 6

Write a Python program to generate all permutations of a list in Python.

In mathematics, the notion of permutation relates to the act of arranging all the members of a set into some sequence or order, or if the set is already ordered, rearranging (reordering) its elements, a process called permuting. These differ from combinations, which are selections of some members of a set where order is disregarded.

<!-- grade_id: cal_days_ans, points: 0, grade: false, solution: true -->
```python
import itertools
def perm_list(lst):
    """Write a Python program to calculate number of days between two dates."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
perm_list([1,2, 4])
```

<!-- grade_id: cal_days_test, points: 1, grade: true, solution: false -->
```python
assert perm_list([1,2]) == [(1, 2), (2, 1)]
assert perm_list([1,2, 4]) == [(1, 2, 4), (1, 4, 2), (2, 1, 4), (2, 4, 1), (4, 1, 2), (4, 2, 1)]
```