---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function to replace all odd numbers in the given array with -7. 
`Hint`: take input as list and inside the function covert it into numpy array then perform the function. Your output should be a list so don't forget to convert numpy array to list

<!-- grade_id: replaceOdd_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def replaceOdd(mylist):
    """Replace all odd numbers in the given array with -7. Input list"""
    
    ### BEGIN SOLUTION
    exercise_1 = np.array(mylist)
    exercise_1[exercise_1 % 2 == 1] = -7
    exercise_1 = exercise_1.tolist()
    return exercise_1
    ### END SOLUTION
```

```python
mylist = [-7, -6, -5, 0, 10, 5, 3, 6, 7, 9]
replaceOdd(mylist)
```

Output:
```
[-7, -6, -7, 0, 10, -7, -7, 6, -7, -7]
```

<!-- grade_id: replaceOdd_test, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
# Test case 1
mylist = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
assert replaceOdd(mylist) == [ 0, -7,  2, -7,  4, -7,  6, -7,  8, -7]

# Test case 2
mylist = [10, 11, 12, 13, 14, 15, 16, 17, 18, 19]
assert replaceOdd(mylist) == [10, -7, 12, -7, 14, -7, 16, -7, 18, -7]

# Test case 3
mylist = [-7, -6, -5, 0, 100, 5, 6, 6, 18, 19]
assert replaceOdd(mylist) == [-7, -6, -7, 0, 100, -7, 6, 6, 18, -7]

```

---

## Question 2:

Write a program, which accept list of range `n` and print all values except the first and last.
`Hint`: Please don't use numpy here
`Example` : input `n = 4` and list will be created `[0, 1, 2, 3]` and output should be: `[1, 2]`

<!-- grade_id: rmv_values_ans, points: 0, grade: false, solution: true -->
```python
def rmv_values(n):
    """accept list of range `n` and print all values except the first and last."""
    ### BEGIN SOLUTION
    arr1 = list(range(n))
#     print(arr1)
    result = arr1[1:-1]
    return result
    ### END SOLUTION
```

```python
rmv_values(4)
```

Output:
```
[1, 2]
```

<!-- grade_id: rmv_values_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert rmv_values(3) == [1]
assert rmv_values(7) == [1, 2, 3, 4, 5]
assert rmv_values(8) == [1, 2, 3, 4, 5, 6]
assert rmv_values(0) == []
```

---
## Question 3: 

 Write a program, which accept 2D array with `nan` and replace all `nan` with `0`.  

<!-- grade_id: remNan_ans, points: 0, grade: false, solution: true -->
```python
import numpy as no
def rem_Nan(A):
    """Write your code here to reverse the string"""
    
    ### BEGIN SOLUTION
    A[np.isnan(A)] = 0
    return A
    ### END SOLUTION
```

```python
arr1 = np.array([[1,2,3],[3,np.nan,3]])
arr2 = np.array([[1,2,3],[3,0,3]])

assert np.allclose(rem_Nan(arr1), arr2)
```

<!-- grade_id: remNan_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
# Test Case 1
arr1 = np.array([[1,2,3],[3,np.nan,3]])
arr2 = np.array([[1,2,3],[3,0,3]])
assert np.allclose(rem_Nan(arr1), arr2)

# Test Case 2
arr3 = np.array([[np.nan,2,np.nan],[3,np.nan,3]])
arr4 = np.array([[0,2,0],[3,0,3]])
assert np.allclose(rem_Nan(arr3), arr4)
```

```python

```