<a href="https://colab.research.google.com/github/kl20879/CE880-lad-class/blob/main/week4/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Kun Liu"
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-4 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

### Task: Complete all the three questions given as follows:

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
!pip install -U scikit-learn
```

Output:
```
Requirement already satisfied: scikit-learn in /usr/local/lib/python3.7/dist-packages (1.0.1)
Requirement already satisfied: joblib>=0.11 in /usr/local/lib/python3.7/dist-packages (from scikit-learn) (1.1.0)
Requirement already satisfied: threadpoolctl>=2.0.0 in /usr/local/lib/python3.7/dist-packages (from scikit-learn) (3.0.0)
Requirement already satisfied: numpy>=1.14.6 in /usr/local/lib/python3.7/dist-packages (from scikit-learn) (1.19.5)
Requirement already satisfied: scipy>=1.1.0 in /usr/local/lib/python3.7/dist-packages (from scikit-learn) (1.4.1)
```

```python
import pandas as pd
import numpy as np
from sklearn.cluster import DBSCAN
from sklearn.cluster import KMeans
from sklearn.datasets import make_moons
from sklearn.metrics import adjusted_rand_score
from sklearn.preprocessing import StandardScaler
from sklearn.metrics import silhouette_score, adjusted_rand_score
from sklearn.decomposition import PCA
scaler = StandardScaler()
```

---
## Question 1: 

Write a code to perform `k-means` algorithm on `scaled_features` and return `silhouette_score` as an output

<!-- grade_id: myScore_kmean_ans, points: 0, grade: false, solution: true -->
```python
# Let's make two moons, where n is the number of samples
features, true_labels = make_moons(n_samples=250, noise=0.05, random_state=42)

# fit_transform() method is basically the combination of fit method and transform method, 
# it is equivalent to fit().transform(). This method performs fit and transform on the 
# input data at a single time and converts the data points. 
# If we use fit and transform separate when we need both then it 
# will decrease the efficiency of the model so we use fit_transform() 
# which will do both the work.
scaled_features = scaler.fit_transform(features)
myscore = []

def myScore_kmean(scaled_features):
    """Write code to perform k-means algorithm and print silhouette_score"""
    cluster = KMeans(n_clusters=2, random_state=0).fit(scaled_features)
    pre = cluster.fit_predict(scaled_features)
    return round(silhouette_score(scaled_features,pre),1)
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: myScore_kmean_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
assert myScore_kmean(scaled_features) == 0.5
```

---
## Question 2: 

Write a code to perform `DBSCAN` algorithm on `scaled_features` and return `silhouette_score` as an output

<!-- grade_id: myScore_DBSCAN_ans, points: 0, grade: false, solution: true -->
```python
def myScore_DBSCAN(scaled_features):
    """Write code to perform DBSCAN algorithm and print silhouette_score"""
    # YOUR CODE HERE
    cluster = DBSCAN().fit(scaled_features)
    pre = cluster.fit_predict(scaled_features)
    return round(silhouette_score(scaled_features,pre),2)
#     raise NotImplementedError()

myScore_DBSCAN(scaled_features)
```

Output:
```
0.38
```

<!-- grade_id: myScore_DBSCAN_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
assert myScore_DBSCAN(scaled_features) == 0.38
```

## Question 3: 

Complete this code perform `PCA`  on `iris data` and return `explained_variance_ratio_` as an output for the first two principal components

<!-- grade_id: myPCA_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
url = "https://archive.ics.uci.edu/ml/machine-learning-databases/iris/iris.data"
# load dataset into Pandas DataFrame
df = pd.read_csv(url, names=['sepal length','sepal width','petal length','petal width','target'])

features = ['sepal length', 'sepal width', 'petal length', 'petal width']
# Separating out the features
x = df.loc[:, features].values
# Separating out the target
y = df.loc[:,['target']].values
# Standardizing the features
x = StandardScaler().fit_transform(x)
    
def myPCA(x):
    """Complete this code return the explained variance ratio for the first two prinical 
    components after doing PCA on x """
    # YOUR CODE HERE
    pca = PCA(n_components=2) #实例化
    pca = pca.fit(x) #拟合模型
    X_dr = pca.transform(x)
    return pca.explained_variance_ratio_
    raise NotImplementedError()

```

<!-- grade_id: myPCA_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
x= myPCA(x)
y = [0.7277045209380135, 0.23030523267680633]
np.testing.assert_allclose(x, y, rtol=1e-5, atol=0)
```

```python

```