<a href="https://colab.research.google.com/github/kl20879/CE880-lad-class/blob/main/week3/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Kun Liu"
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-3 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function to calculate the median  
`Hint`: To find the median, you first need to sort the values in your sample. The `median` of a sample of numeric data is the value that lies in the middle when we sort the data. The data may be sorted in ascending or descending order, the median remains the same.

<!-- grade_id: my_median_ans, points: 0, grade: false, solution: true -->
```python
def my_median(sample):
    """Function for calculating median"""
    # YOUR CODE HERE
    median = 0
    Lcopy = sorted(sample)
    if len(Lcopy)%2==1:
      median = Lcopy[len(Lcopy)//2]
    else:
      median = (Lcopy[len(Lcopy)//2]+Lcopy[len(Lcopy)//2-1])/2
    return median
    raise NotImplementedError()
```

```python
my_median([3, 5, 1, 4, 2, 6])
```

Output:
```
3.5
```

<!-- grade_id: my_median_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
# Test case 1
assert my_median([3, 5, 1, 4, 2]) == 3
assert my_median([1, 4, 2, 5, 0]) == 2
assert my_median([1, 4, 2, 5, 0, 5, 6, 3, 5]) == 4
assert my_median([3, 5, 1, 4, 2, 6]) == 3.5
```

---

## Question 2:

Caluate Interquartile Range of the given list.

`Please read me before attemping the question`
A quartile is a type of quantile. The first quartile `(Q1)`, is defined as the middle number between the smallest number and the median of the data set, the second quartile `(Q2)` – median of the given data set while the third quartile `(Q3)`, is the middle number between the median and the largest value of the data set.

Algorithm to find Quartiles :
Quartiles are calculated by the help of the median. If the number of entries is an even number i.e. of the form 2n, then, first quartile `(Q1)` is equal to the median of the `n` smallest entries and the third quartile `(Q3)` is equal to the median of the `n` largest entries.

If the number of entries is an odd number i.e. of the form `(2n + 1)`, then

* the first quartile (Q1) is equal to the median of the n smallest entries
* the third quartile (Q3) is equal to the median of the n largest entries
* the second quartile(Q2) is the same as the ordinary median.

Example: 
input = [75, 69, 56, 46, 47, 79, 92, 97, 89, 88, 36, 96, 105, 32, 116, 101, 79, 93, 91, 112]
After sorting the above data set:
[32, 36, 46, 47, 56, 69, 75, 79, 79, 88, 89, 91, 92, 93, 96, 97, 101, 105, 112, 116]
Here the total number of terms is 20.

The second quartile `(Q2)` or the median of the above data is `(88 + 89) / 2 = 88.5`
The first quartile `(Q1)` is median of first `n` i.e. 10 terms (or `n` i.e. 10 smallest values) = 62.5
The third quartile `(Q3)` is the median of `n` i.e. 10 largest values (or last `n` i.e. 10 values) = 96.5
Then, `IQR = Q3 – Q1 = 96.5 – 62.5 = 34.0`

`hint`: Please explore `np.percentile` function

<!-- grade_id: myfibonacci_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def my_quartile(data):
    """Caluate Interquartile Range of the given list."""
    # YOUR CODE HERE
    Le = len(data)
    if Le%2 == 0:
      n = len(data)//2
      Q2 = my_median(data)
      Lcopy = sorted(data)
      Q1 = my_median(Lcopy[:n])
      LLcopy = sorted(data,reverse=1)
      Q3 = my_median(LLcopy[:n])
      IQR = Q3-Q1
    return IQR  
    raise NotImplementedError()
```

```python
data = [13, 26, 26, 17, 16, 49, 37, 29, 40, 13, 12, 23, 0]
my_quartile(data)
```

Output:
```
23
```

<!-- grade_id: myfibonacci_test, points: 1, grade: true, solution: false -->
```python
"""Check that correct answers here"""
data = [32, 36, 46, 47, 56, 69, 75, 79, 79, 88, 89, 91, 92, 93, 96, 97, 
        101, 105, 112, 116]
assert my_quartile(data) == 34

data = [3, 6, 6, 7, 6, 9, 17, 19, 10, 11, 22, 13, 5]
assert my_quartile(data) == 7

data = [13, 26, 26, 17, 16, 49, 37, 29, 40, 13, 12, 23, 0]
assert my_quartile(data) == 16
```

---
## Question 3: 

Write a program to calculate the `standard deviation` of a dataset. Please create a function that first calculate the variance and calculate `standard deviation`. Please do not use inbuilt `numpy.var` and `numpy.std` functions. Write your own codes. 

<!-- grade_id: my_std_ans, points: 0, grade: false, solution: true -->
```python
import math

def my_std(data, ddof=0):
    """Write a program to calculate the standard deviation of a dataset."""
    
    # YOUR CODE HERE
    u = sum(data)/len(data)
    sumall = 0
    for num in data:
      sumall +=(num-u)**2
    sumall = sumall/len(data)
    res = math.sqrt(sumall)
    return res
    raise NotImplementedError()
```

```python
data = [4, 8, 6, 5, 3, 2, 8, 9, 2, 5]
my_std(data)
```

Output:
```
2.4
```

<!-- grade_id: my_std_test, points: 1, grade: true, solution: false -->
```python
"""Test you solution"""
# Test Case 1
data = [4, 8, 6, 5, 3, 2, 8, 9, 2, 5]
assert my_std(data) == 2.4

data = [24, 18, 26, 15, 13, 10, 11, 4, 4, 1]
assert math.isclose(my_std(data), 7.97, rel_tol=0.05)

data = [14, 18, 16, 15, 13, 12, 18, 19, 12, 15]
assert math.isclose(my_std(data), 2.4, rel_tol=0.05)
```