<a href="https://colab.research.google.com/github/kl20879/CE880-lad-class/blob/main/week3/problem2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Kun Liu"
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-3 : Problem 2 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

```python
# Please make sure if pandas is installed and if not then run the this cell otherwise skip
!pip install pandas
```

Output:
```
Requirement already satisfied: pandas in /usr/local/lib/python3.7/dist-packages (1.1.5)
Requirement already satisfied: pytz>=2017.2 in /usr/local/lib/python3.7/dist-packages (from pandas) (2018.9)
Requirement already satisfied: python-dateutil>=2.7.3 in /usr/local/lib/python3.7/dist-packages (from pandas) (2.8.2)
Requirement already satisfied: numpy>=1.15.4 in /usr/local/lib/python3.7/dist-packages (from pandas) (1.19.5)
Requirement already satisfied: six>=1.5 in /usr/local/lib/python3.7/dist-packages (from python-dateutil>=2.7.3->pandas) (1.15.0)
```

#### Let's load iris dataset

```python
import pandas as pd
import numpy as np
df = pd.read_csv('iris.csv') # Please adjust the path is needed if you are loading the datafile from local computer
df # display the DataFrame
```

---
## Question 4:

Write a program to find the mean of the `petal.length` variable in the iris data

<!-- grade_id: filter_1_ans, points: 0, grade: false, solution: true -->
```python
def iris_mean():
    """write a function to find the mean of the `petal.length`."""
    
    # YOUR CODE HERE
    res = df['petal.length'].mean()
    return res
    raise NotImplementedError()
```

<!-- grade_id: filter_1_test, points: 1, grade: true, solution: false -->
```python
assert np.isclose(iris_mean(), 3.75, atol= 0.1)

```

---
## Question 5:
Write a program to count the number of rows in the iris data, where `variety = Setosa`

<!-- grade_id: iris_count_ans, points: 0, grade: false, solution: true -->
```python
def iris_count():
    """Write a function to count the number of rows, where variety is Setosa"""
    
    # YOUR CODE HERE
    res = len(df[df['variety']=='Setosa'])
    return res
    raise NotImplementedError()
```

<!-- grade_id: iris_count_test, points: 1, grade: true, solution: false -->
```python
"""Check the your answer"""
assert iris_count() == 50
```

## Question 6

Write a Pandas program to calculate the correlation between `sepal.length` and `sepal.width`.

<!-- grade_id: iris_correlation_ans, points: 0, grade: false, solution: true -->
```python
import itertools
def iris_correlation():
    """calculate the correlation"""
    # YOUR CODE HERE
    res = pd.DataFrame(df['sepal.length']).corrwith(df['sepal.width'])
    return res
    raise NotImplementedError()
```

```python
iris_correlation()
```

<!-- grade_id: iris_correlation_test, points: 1, grade: true, solution: false -->
```python
assert np.isclose(iris_correlation(), -0.11, atol= 0.05)
```

```python

```