<a href="https://colab.research.google.com/github/kl20879/CE880-lad-class/blob/main/week5/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Kun Liu"
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-5 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
!pip install -U scikit-learn
```

Output:
```
Requirement already satisfied: scikit-learn in /usr/local/lib/python3.7/dist-packages (1.0.1)
Requirement already satisfied: scipy>=1.1.0 in /usr/local/lib/python3.7/dist-packages (from scikit-learn) (1.4.1)
Requirement already satisfied: threadpoolctl>=2.0.0 in /usr/local/lib/python3.7/dist-packages (from scikit-learn) (3.0.0)
Requirement already satisfied: joblib>=0.11 in /usr/local/lib/python3.7/dist-packages (from scikit-learn) (1.1.0)
Requirement already satisfied: numpy>=1.14.6 in /usr/local/lib/python3.7/dist-packages (from scikit-learn) (1.19.5)
```

```python
import sys
import warnings

if not sys.warnoptions:
    warnings.simplefilter("ignore")
```

---
## Question 1: 
We’re ready to split a larger dataset to solve a regression problem. You’ll use a well-known Boston house prices dataset, which is included in sklearn. This dataset has 506 samples, 13 input variables, and the house values as the output. We have loaded dataset for you and splited into training and test set. Write code to fit linear regression on training `input (x)` and training `output (y)` and print the training score using `model.score` function. 

```python
from sklearn.datasets import load_boston
from sklearn.model_selection import train_test_split
x, y = load_boston(return_X_y=True)
x_train, x_test, y_train, y_test = train_test_split(x, y, test_size=0.4, random_state=0)
```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing dataset (i.e.
```

<!-- grade_id: my_regression_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

def my_regression(x_train, y_train):
    """Write code to fit linear regression on training input (x) and training output (y) and 
    print the training score using model.score function. """
    # YOUR CODE HERE
    lr=LinearRegression()
    lr.fit(x_train, y_train)
    return lr.score(x_train, y_train)
    raise NotImplementedError()
```

<!-- grade_id: my_regression_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_regression(x_train, y_train), 0.7668, rel_tol = 0.05)
```

---
## Question 2: 

Now, implement `RandomForestRegressor` on the same dataset and check the performance on the `test` dataset. Hint: Please read the question carefuly. 

<!-- grade_id: my_RMregression_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestRegressor

def my_RMregression(x_train, y_train, x_test, y_test):
    """Write code to fit RandomForestRegressor on training input (x) and training output (y) and 
    print the test score using model.score function. """
    # YOUR CODE HERE
    rf=RandomForestRegressor()
    rf.fit(x_train, y_train)
    return rf.score(x_test, y_test)
    raise NotImplementedError()
```

<!-- grade_id: my_RMregression_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_RMregression(x_train, y_train, x_test, y_test), 0.8325, rel_tol = 0.05)
```

---
## Question 3: 

Write code to fit classificaiton model LogisticRegression on training `input (x)` and training `output (y)` and 
print the predicted (x) that is your same data on which you have trained your model 

<!-- grade_id: my_classification_ans, points: 0, grade: false, solution: true -->
```python

import numpy as np
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import classification_report, confusion_matrix
x = np.arange(10).reshape(-1, 1)
y = np.array([0, 0, 0, 0, 1, 1, 1, 1, 1, 1])

def my_classification(x, y):
    """Write code to fit classificaiton model LogisticRegression on training input (x) and training output (y) and 
    print the predicted (x) that is your same data on which you have trained your model """
    # YOUR CODE HERE
    lr=LogisticRegression(max_iter=5,C=0.1)
    lr.fit(x, y)
    return lr.predict(x)
    raise NotImplementedError()

```

<!-- grade_id: my_classification_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
y_hat = my_classification(x, y)
y_hat = y_hat.tolist()
assert y_hat == [0, 0, 0, 1, 1, 1, 1, 1, 1, 1]
```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/linear_model/_logistic.py:818: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  extra_warning_msg=_LOGISTIC_SOLVER_CONVERGENCE_MSG,
```

```python

```

```python

```